# Lab 08: Circular, Array-backed Queue

## Overview

For this assignment you will implement a circular, array-backed queue data structure.

In the following class, which you are to complete, the backing array will be created and populated with `None`s in the `__init__` method, and the `head` and `tail` indexes set to sentinel values (you shouldn't need to modify `__init__`). Enqueueing and Dequeueing items will take place at the tail and head, with `tail` and `head` tracking the position of the most recently enqueued item and that of the next item to dequeue, respectively. To simplify testing, your implementation should make sure that when dequeuing an item its slot in the array is reset to `None`, and when the queue is emptied its `head` and `tail` attributes should be set to `-1`.

Because of the fixed size backing array, the `enqueue` operation is defined to raise a `RuntimeError` when the queue is full — the same exception should be raised when `dequeue` is called on an empty queue.

Finally, the `resize` method will allow the array underlying the queue to be increased in size. It is up to you how to implement this (you can either leave the elements in their current positions, though this may require "unwrapping" elements, or you can simply move all elements towards the front of the array). You may assume that `resize` will only be called with a value greater than the current length of the underlying array.

```python
def dequeue(self):
        if self.empty():
            raise RuntimeError
        if self.head == self.tail:
            out = self.data[self.head]
            self.head = -1
            self.tail = -1
            return out
        if self.head == len(self.data):
            self.head = 0
            out = self.data[self.head]
            self.data[self.head] = None
            self.head += 1
        else:
            out = self.data[self.head]
            self.data[self.head] = None
            self.head += 1
        return out
```

<!-- grade_id: queue, points: 0, grade: false, solution: true -->
```python
class Queue:
    def __init__(self, limit=10):
        self.data = [None] * limit
        self.head = -1
        self.tail = -1

    def enqueue(self, val):
        for x in self.data:
            if x == None:
                break
        else:
            raise RuntimeError
        if self.head == -1 and self.tail == -1:
            self.head = 0
            self.tail = 0
            self.data[self.tail] = val
        elif self.tail == len(self.data) - 1 and self.head != 0:
            self.tail = 0
            self.data[self.tail] = val
        else:
            self.tail += 1
            self.data[self.tail] = val
        
    def dequeue(self):
        if self.empty():
            raise RuntimeError
        if self.head == self.tail:
            out = self.data[self.head]
            self.data[self.head] = None
            self.head = -1
            self.tail = -1
            return out
        if self.head == len(self.data):
            self.head = 0
            out = self.data[self.head]
            self.data[self.head] = None
            self.head += 1
        else:
            out = self.data[self.head]
            self.data[self.head] = None
            self.head += 1
        return out
    
    def resize(self, newsize):
        assert(len(self.data) < newsize)
        size = len(self.data)
        out = [None] * newsize
        idx = self.head
        nidx = 0
        count = 0
        while count != size:
            if idx != len(self.data):
                out[nidx] = self.data[idx]
                nidx += 1
            else:
                idx = 0
                out[nidx] = self.data[idx]
                nidx += 1
            count += 1
            idx += 1
        count = 0
        for x in out:
            if x != None:
                count += 1
        self.head = 0
        self.tail = size - 1
        self.data = out
        
    def empty(self):
        for x in self.data:
            if x != None:
                return False
        return True
    
    def __bool__(self):
        return not self.empty()
    
    def __str__(self):
        if not(self):
            return ''
        return ', '.join(str(x) for x in self)
    
    def __repr__(self):
        return str(self)
    
    def __iter__(self):
        idx = self.head
        for i in range(len(self.data)):
            if self.data[idx] != None and idx != len(self.data):
                yield self.data[idx]
            idx = idx + 1
            if idx == len(self.data):
                idx = 0
            
```

<!-- grade_id: queue_tests_1, points: 5, grade: true, solution: false -->
```python
# (5 points)

from unittest import TestCase
tc = TestCase()

q = Queue(5)
tc.assertEqual(q.data, [None] * 5)

for i in range(5):
    q.enqueue(i)
    
with tc.assertRaises(RuntimeError):
    q.enqueue(5)

for i in range(5):
    tc.assertEqual(q.dequeue(), i)
    
tc.assertTrue(q.empty())
```

<!-- grade_id: queue_tests_2, points: 5, grade: true, solution: false -->
```python
# (5 points)

from unittest import TestCase
tc = TestCase()

q = Queue(10)

for i in range(6):
    q.enqueue(i)
    
tc.assertEqual(q.data.count(None), 4)

for i in range(5):
    q.dequeue()
    
tc.assertFalse(q.empty())
tc.assertEqual(q.data.count(None), 9)
tc.assertEqual(q.head, q.tail)
tc.assertEqual(q.head, 5)

for i in range(9):
    q.enqueue(i)

with tc.assertRaises(RuntimeError):
    q.enqueue(10)

for x, y in zip(q, [5] + list(range(9))):
    tc.assertEqual(x, y)
    
tc.assertEqual(q.dequeue(), 5)
for i in range(9):
    tc.assertEqual(q.dequeue(), i)

tc.assertTrue(q.empty())
```

<!-- grade_id: queue_tests_3, points: 5, grade: true, solution: false -->
```python
# (5 points)

from unittest import TestCase
tc = TestCase()

q = Queue(5)
for i in range(5):
    q.enqueue(i)
for i in range(4):
    q.dequeue()
for i in range(5, 9):
    q.enqueue(i)
    
with tc.assertRaises(RuntimeError):
    q.enqueue(10)

q.resize(10)

for x, y in zip(q, range(4, 9)):
    tc.assertEqual(x, y)
    
for i in range(9, 14):
    q.enqueue(i)

for i in range(4, 14):
    tc.assertEqual(q.dequeue(), i)
    
tc.assertTrue(q.empty())
tc.assertEqual(q.head, -1)
```